# Assignment Data Science

## Notebook made by

**Gebruik graag dit formaat**

* Voor de namen:  voornaam rest van je naam, voornaam rest van je naam,....
* je studentnummers: hetzelfde: scheidt met `,`
* je emails: hetzelfde: scheidt met `,`
* voor je groep: alleen de hoofdletter (dus voor de groep van Marx zou je `A` kiezen)

__Groepsnaam__: TeamEenmansleger

__Namen__:Anoniem
__Emails__:Anoniem
__Student id__:Anoniem
__Groep__ : B

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar zichtbare tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. De andere tests worden na inleveren toegevoegd aan die cell.

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Week-1" data-toc-modified-id="Week-1-1"><span class="toc-item-num">1&nbsp;&nbsp;</span>Week 1</a></span></li><li><span><a href="#Deel-1-Quizmaster-paradox-(2pt)" data-toc-modified-id="Deel-1-Quizmaster-paradox-(2pt)-2"><span class="toc-item-num">2&nbsp;&nbsp;</span>Deel 1 Quizmaster paradox (2pt)</a></span><ul class="toc-item"><li><span><a href="#Verschil-tussen-theoretische-kans-en-de-experimentele-kans-om-te-winnen." data-toc-modified-id="Verschil-tussen-theoretische-kans-en-de-experimentele-kans-om-te-winnen.-2.1"><span class="toc-item-num">2.1&nbsp;&nbsp;</span>Verschil tussen theoretische kans en de experimentele kans om te winnen.</a></span></li><li><span><a href="#significant?" data-toc-modified-id="significant?-2.2"><span class="toc-item-num">2.2&nbsp;&nbsp;</span>significant?</a></span></li></ul></li><li><span><a href="#Deel-2:-Log-normal-distribution-of-Wikipedia" data-toc-modified-id="Deel-2:-Log-normal-distribution-of-Wikipedia-3"><span class="toc-item-num">3&nbsp;&nbsp;</span>Deel 2: Log normal distribution of Wikipedia</a></span><ul class="toc-item"><li><span><a href="#Vraag-W1--(5pts)" data-toc-modified-id="Vraag-W1--(5pts)-3.1"><span class="toc-item-num">3.1&nbsp;&nbsp;</span>Vraag W1  (5pts)</a></span></li><li><span><a href="#Vraag-W2-(1pt)" data-toc-modified-id="Vraag-W2-(1pt)-3.2"><span class="toc-item-num">3.2&nbsp;&nbsp;</span>Vraag W2 (1pt)</a></span></li></ul></li><li><span><a href="#Deel-3:-theorie" data-toc-modified-id="Deel-3:-theorie-4"><span class="toc-item-num">4&nbsp;&nbsp;</span>Deel 3: theorie</a></span><ul class="toc-item"><li><span><a href="#Variant" data-toc-modified-id="Variant-4.1"><span class="toc-item-num">4.1&nbsp;&nbsp;</span>Variant</a></span></li></ul></li><li><span><a href="#Deel-4-Corona:-Wat-is-erger:-besmettelijkheid-of-dodelijkheid?" data-toc-modified-id="Deel-4-Corona:-Wat-is-erger:-besmettelijkheid-of-dodelijkheid?-5"><span class="toc-item-num">5&nbsp;&nbsp;</span>Deel 4 Corona: Wat is erger: besmettelijkheid of dodelijkheid?</a></span><ul class="toc-item"><li><span><a href="#Klopt-het-wel?-(1pt)" data-toc-modified-id="Klopt-het-wel?-(1pt)-5.1"><span class="toc-item-num">5.1&nbsp;&nbsp;</span>Klopt het wel? (1pt)</a></span></li><li><span><a href="#Aangepast-dataframe-(1pt)" data-toc-modified-id="Aangepast-dataframe-(1pt)-5.2"><span class="toc-item-num">5.2&nbsp;&nbsp;</span>Aangepast dataframe (1pt)</a></span></li></ul></li></ul></div>

# Week 1

* Elke vraag is 1 punt waard, tenzij anders aangegeven.

```python
import bz2
import codecs
import os
import random
import re
from collections import Counter, defaultdict

import matplotlib.pyplot as plt
import numpy as np
import pandas as pd
import seaborn as sn
from bs4 import BeautifulSoup
from nose.tools import assert_almost_equal, assert_equal
from tqdm import notebook

%matplotlib inline
```

# Deel 1 Quizmaster paradox (2pt)


Hier is een leuk en misschien verrassend statistisch raadsel, en een goede manier om wat te oefenen met het schrijven van python-functies.

In een gameshow proberen deelnemers te raden welke van de 3 gesloten deuren een geldprijs bevatten (geiten zitten achter de andere twee deuren). Natuurlijk, de kans om de juiste deur te kiezen is 1 op 3. Als een twist opent de gastheer van de show af en toe een deur nadat een deelnemer zijn of haar keuze heeft gemaakt. Deze deur is altijd een van de twee die de deelnemer niet heeft gekozen, en is ook altijd een van de geitendeuren (merk op dat het altijd mogelijk is om dit te doen, aangezien er twee geitendeuren zijn). Op dit punt heeft de deelnemer de mogelijkheid om zijn of haar oorspronkelijke keuze te behouden, of naar de andere ongeopende deur te switchen. De vraag is: is er een voordeel bij het wisselen van deuren? Het antwoord verrast veel mensen die de vraag nog niet eerder hebben gehoord.

We kunnen het probleem oplossen door simulaties uit te voeren in Python. 
In onderstaande 2 cellen wordt dat gedaan. Het `experiment` deel is al geprogrammeerd. Jij moet de functie `speel` afmaken. Het eerste stukje hebben we al voor je geprogrammeerd.  Die moet dus werken voor elk aantal deuren (van minstens 3), en ook in 2 varianten.

**HINTS**
1. Gebruik set theory operaties (als `-,|,&`) om handig deuren uit te sluiten of toe te voegen.
2. Je kunt met `np.random.choice` ook meerdere items uit een lijst tegelijk pakken.
3. Zorg dat je het probleem begrijpt. Het is echt lastig hoor! Maar er over denken met meerdere deuren en met de twee manieren van openen (bij meer deuren dan 3 speelt dat) helpt je het te snappen.
4. Bij mijn implementatie kreeg ik pas een mooie grafiek te zien bij N=10.000. Maar dat duurt net wat langer. De trend is al wel heel goed zichtbaar bij N=100.

<!-- grade_id: qmb, points: 0, grade: false, solution: true -->
```python
def speel(aantaldeuren=3, soort='eentje'):
    '''Return True als je de prijs wint bij wisselen; anders False.
    Bij soort=eentje openen we 1 deur. Anders openen we ze allemaal behalve 2.
    We openen nooit de gekozen deur of de deur waar de prijs achter zit.'''
    deuren = set(np.arange(aantaldeuren))
    prijs_deur = np.random.choice(list(deuren))
    gekozen_deur = np.random.choice(list(deuren))
    
    if soort == 'eentje':    
        # quizmaster opent deur, en die valt af uit nieuwe keuzes voor deelnemer
        afgevallen_deur = np.random.choice(list(deuren - {prijs_deur, gekozen_deur}))
        
        # nieuwe deur wordt gekozen voor deelnemer
        gekozen_deur2 = np.random.choice(list(deuren - {gekozen_deur, afgevallen_deur}))
        
        # check voor winst na wisselen van deur
        return gekozen_deur2 == prijs_deur

    else:
        # deelnemer wint alleen als hij niet de winnende deur had gekozen in eerste instantie
        return prijs_deur != gekozen_deur


def experiment(aantaldeuren=3, soort='eentje', N=10**4):
    '''Run N keer het experiment van type soort met aantaldeuren, en geef voor elk aantal deuren, de gemiddelde win kans.'''
    return np.array([speel(aantaldeuren, soort) for _ in range(N)]).mean()
```

```python
# HIERMEE RUNNEN WE N KEER VOOR ELK AANTAL DEUREN TUSSEN 3 EN aantaldeuren HET EXPEREIMENT, STEEDS
# VOOR HET OPENEN VAN 1 EN VAN ALLE BEHALVE 1 DEUR.
def data(N=1000, aantaldeuren=21):
    '''Run het experiment voor elk aantal deuren tussen 3 en aantaldeuren N keer'''
    uitkomsten = {
        D: {
            'eentje_open': experiment(D, 'eentje', N),
            'alle_behalve_eentje': experiment(D, 'alle_behalve_eentje', N)
        }
        for D in range(3, aantaldeuren)
    }
    return uitkomsten


# NU MAKEN WE ER EEN PANDAS DATAFRAME VAN EN PLOTTEN HET.
# BEKIJK ZELF df EENS IN EEN NIEUWE CEL.
df = pd.DataFrame(data(10**4)).T
df.index.name = 'aantal deuren'
Titel = 'Kans om de prijs te winnen bij switchen bij N deuren.'
df.plot(kind='bar', title=Titel);
```

<!-- grade_id: qmt, points: 1, grade: true, solution: false -->
```python
df = pd.DataFrame(data(10**3)).T
assert df.shape == (18, 2)
assert set(df.columns) == {'alle_behalve_eentje', 'eentje_open'}
```

## Verschil tussen theoretische kans en de experimentele kans om te winnen.

Nu snap je vast het probleem heel goed (.....). Vooral na het laatste plaatje. Nou, dan kan je ook de theoretische kansen om te winnen na switchen voor N deuren wel berekenen toch?

Doe dat dan maar. Weer voor de twee gevallen.

We vragen om 2 arrays van lengte 18 waarin je voor elk aantal deuren in `range(3,21)` de kans berekent dat je wint na switchen. Voor de beide gevallen. 

Bij het nakijken van de vorige vraag gaan we jouw experimentele antwoorden vergelijken met de juiste theoretische kansen. Bij N=1000 verwachten we per deur alleen verschillen kleiner dan 5%. Heb je die zelf ook?

<!-- grade_id: Vraag1, points: 0, grade: false, solution: true -->
```python
# kans dat winnende deur wordt geopend:
# (alle deuren - gekozen deur) / alle deuren / (alle deuren - winnende deur - gekozen deur)
Kans_eentje_open = np.array([((N - 1) / N / (N - 2)) for N in range(3, 21)])

# kans dat winnende deur wordt geopend:
# (alle deuren - gekozen deur) / alle deuren
Kans_alle_behalve_eentje_open = np.array([((N - 1) / N) for N in range(3, 21)])


# Nu berekenen we de verschillen in uitkomsten tussen die we vinden in jouw experiment (dat staat in df) 
# en jouw theoretusche kans.
# Bij mij was het gemiddelde verschil over al die deuren 0.008 en 0.007
Errors = {
    'alle_behalve_eentje':
    abs(df.alle_behalve_eentje.values - Kans_alle_behalve_eentje_open),
    'eentje_open':
    abs(df.eentje_open.values - Kans_eentje_open)
}

# Maak een dataframe van die dict
Errors = pd.DataFrame(Errors)
Errors.index = df.index

print(Errors.describe())
Errors.plot(
    kind='bar',
    title=
    'Verschil tussen theoretische kans en de experimentele kans om te winnen.')
plt.show()
```

<!-- grade_id: Vraag1_antwoord, points: 1, grade: true, solution: false -->
```python
assert len(Kans_eentje_open) == len(Kans_alle_behalve_eentje_open) == 18

```

## significant?

Laten we hopen dat het verschil tussen jouw theoretische kans en je experimentele kans niet significant is, want dan is er vast iets mis met 1 van beide. 

Nou dat is een mooie test op jezelf, dan toch? 

Dus doe een student T-test, en beschrijf wat de uitkomst betekent. Moet ie nou one- of two-sided?

<!-- grade_id: Vraag1aa, points: 1, grade: true, solution: true -->
```python
from scipy import stats

# maak dict met uitkomsten
Compare = {
    'daadwerkelijke_alle_behalve_eentje': df.alle_behalve_eentje.values, 
    'kans_alle_behalve_eentje': Kans_alle_behalve_eentje_open,
    'daadwerkelijke_eentje_open': df.eentje_open.values, 
    'kans_eentje_open': Kans_eentje_open
}

# maak een dataframe van die dict
Compare = pd.DataFrame(Compare)
Compare.index = df.index

# voer t-tests uit op beide versies
ttest_alles_behalve_eentje = stats.ttest_ind(Compare['daadwerkelijke_alle_behalve_eentje'], 
                                             Compare['kans_alle_behalve_eentje'], equal_var=True)
ttest_eentje_open = stats.ttest_ind(Compare['daadwerkelijke_eentje_open'], 
                                    Compare['kans_eentje_open'], equal_var=True)

display(Compare.head())
ttest_alles_behalve_eentje, ttest_eentje_open

# both p values are > .05 which means that there is no significant different in groups
```

# Deel 2: Log normal distribution of Wikipedia

1. Lees het artikel van Jake van der Plas over de ["line lengths in Python modules"](http://jakevdp.github.io/blog/2017/11/09/exploring-line-lengths-in-python-packages)
2. In dit deel gaan we iets vergelijkbaars doen, alleen dan met de Nederlandse Wikipedia.

4. We are interested in the following things:
    1. distribution of line lengths on the Dutch Wikipedia (measured in number of characters)
    2. distribution of article (called pages in the dump) lengths on the Dutch Wikipedia (measured in number of lines)
    3. distribution of article (called pages in the dump) lengths on the Dutch Wikipedia (measured  in number of characters)
5. Your task is to find out if they follow a lognormal distribution.

```python
def loadfile():
    if 'Wikipedia.bz2' in os.listdir():
        return 'Wikipedia.bz2'
    elif os.path.exists('../../data/Week1/'):
        return '../../data/Week1/Wikipedia.bz2'
    elif os.path.exists('../../../data/Week1/'):
        return '../../../data/Week1/Wikipedia.bz2'
```

## Vraag W1  (5pts)

Vul de code in de cell aan zodat je voor de file de volgende metrieken krijgt:
* `char_list`, een lijst met lengtes van de alle Wikipediapages, uitgedrukt in het aantal karakters 
    * Dus als er `n` Wikipediapaginas zijn, heeft deze lijst lengte `n`.
* `lines_list`, een lijst met lengtes van  alle Wikipediapages, uitgedrukt in het aantal regels
    * Dus als er `n` Wikipediapaginas zijn, heeft deze lijst lengte `n`.
* `lines_lens`, een lijst met de lengtes van alle regels op Wikipedia, uitgedrukt in het aantal karakters 
    * de lengte van deze lijst is dus de som van alle elementen in de vorige lijst.

De functie `wikipedia_distribution()` bevat ook de variabele `clean`, die True of False is. Bij `True` moeten irrelevant karakters, zoals HTML weggehaald worden.
    
Opmerkingen:
* Het is belangrijk dat je efficiënte code schrijft. Het is een groot bestand, en je wil dat het niet te lang duurt.
* Negeer regels die na cleanen 0 karakters over hebben.
* Een regel (*line*) is alle tekst tot aan een `\n` karakter. Dus echt een regel zoals op een typemachine. Als je een file regel voor regel leest heb je precies de `line`s die we hier bedoelen.

<!-- grade_id: cell-930934950b6b5842, points: 0, grade: false, solution: true -->
```python
# voor het cleanen (enkel HTML tags verwijderen)
CLEANR = re.compile('<.*?>') 

def wikipedia_distribution(file, clean=True):
    with bz2.open(file, encoding="utf8", mode='rt') as source_file:        
        # creëer de lijsten
        char_list = []
        lines_list = []
        lines_lens = []
        
        # staat: laatst gezien is einde pagina (om nieuwe te kunnen tracen)
        state = 0
        
        # zet de tellers per wikipagina op nul
        chars_per_page = 0
        lines_per_page = 0
        
        # itereer over regels
        for line in notebook.tqdm(source_file):          
            # begin pagina gevonden
            if '<page>' in line:
                # en laatste was pagina einde
                if state == 0:
                    # dan page begin gevonden
                    state = 1
                # als laatste was pagina begin -> onvolledige pagina
                else:
                    # reset paginalengte
                    chars_per_page = 0
                    lines_per_page = 0
            
            # als in pagina zit
            if state == 1:
                # als er schoongemaakt moet worden
                if clean:
                    # voeg schoongemaakte lengte toe
                    chars_per_page += len(re.sub(CLEANR, '', line))
                    lines_per_page += 1
                    lines_lens.append(len(re.sub(CLEANR, '', line)))
                else:
                    # anders voeg lengte toe
                    chars_per_page += len(line)
                    lines_per_page += 1
                    lines_lens.append(len(line))
                
            # einde pagina gevonden
            if '</page>' in line:
                # en laatste was pagina begin
                if state == 1:
                    # voeg totale lengte toe aan list
                    char_list.append(chars_per_page)
                    lines_list.append(lines_per_page)
                # reset paginalengte
                chars_per_page = 0
                lines_per_page = 0
                # eind pagina gevonden
                state = 0
                
        return char_list, lines_list, lines_lens

# maak de lijsten
lijsten = wikipedia_distribution(loadfile(), clean=True)

# bereken sommen en lengtes
sommen = [sum(lijst) for lijst in lijsten]
lengtes = [len(lijst) for lijst in lijsten]

# maak er een dataframe van voor de duidelijkheid
pd.DataFrame(data = {"sum": sommen, "lenght": lengtes}, 
             columns = ["sum", "lenght"], 
             index = ["char_list", "lines_list", "lines_lens"])
```

Eerst even kijken naar het resultaat zonder cleaning

<!-- grade_id: Vraag10_antwoordA, points: 1, grade: true, solution: false -->
```python
char_list, lines_list, lines_lens = wikipedia_distribution(loadfile(), clean=False)

assert_equal(type(char_list), list)
assert_equal(type(lines_list), list)
assert_equal(type(lines_lens), list)

f, axes = plt.subplots(1, 3, figsize=(20, 5))

sn.histplot(np.log10(char_list), kde=True, ax=axes[0])
axes[0].set_title('Number of characters per page (log)')
sn.histplot(np.log10(lines_list), kde=True, ax=axes[1])
axes[1].set_title('Number of lines per page (log)')
sn.histplot(np.log10(lines_lens), kde=True, ax=axes[2])
axes[2].set_title('Linelengths (log)');

```

<!-- grade_id: Vraag10_AntwoordB, points: 1, grade: true, solution: false -->
```python
# testcell
```

En dan het resultaat met cleaning.

<!-- grade_id: Vraag10_AntwoordC, points: 1, grade: true, solution: false -->
```python
char_list, lines_list, lines_lens = wikipedia_distribution(loadfile(), clean=True)

assert_equal(type(char_list), list)
assert_equal(type(lines_list), list)
assert_equal(type(lines_lens), list)

f, axes = plt.subplots(1, 3, figsize=(20, 5))

sn.histplot(np.log10(char_list), kde=True, ax=axes[0])
axes[0].set_title('Number of characters per page (log)')
sn.histplot(np.log10(lines_list), kde=True, ax=axes[1])
axes[1].set_title('Number of lines per page (log)')
sn.histplot(np.log10(lines_lens), kde=True, ax=axes[2])
axes[2].set_title('Linelengths (log)');
```

<!-- grade_id: cell-c5c0c02f3c8be674, points: 1, grade: true, solution: false -->
```python
# testcell
```

<!-- grade_id: Vraag10_AntwoordD, points: 1, grade: true, solution: false -->
```python

```

## Vraag W2 (1pt)
Beantwoord de volgende vragen in onderstaande markdown cell:

1. Did I remove enough/too much?

2. Above are distribution plots.  Explain what we see, and whether the log-normal hypothesis seems true or not.

3. (For yourself) Relate your findings to those of Jake van der Plas and possibly earlier work on lengths in statistical linguistics.


**Curve fitting**

Je kan de techniek die beschreven wordt in <https://towardsdatascience.com/basic-curve-fitting-of-scientific-data-with-python-9592244a2509> gebruiken om te kijken of je een Gaussian curve kunt fitten op de log waardes, en hoe goed dat fit. 

<!-- grade_id: cell-e0ed1a05dce7585d, points: 1, grade: true, solution: true -->
1. Misschien heb ik niet genoeg weggehaald, maar ik zou niet weten wat er verder nog moet worden weggehaald om het verder op te schonen. Alle HTML tags worden in ieder geval verwijderd. 

2. We zien distributiefiguren waarin het aantal karakters per pagina, het aantal regels per pagina en de lengtes van de regels van Nederlandse Wikipediapagina's worden afgebeeld. De eerste twee lijken aardig overeen te komen met een log-normale distributie, afgezien van de piek aan het begin. Dit kan verklaard worden doordat er grofweg twee soorten wikipediapagina's zijn: korte van slechts een paar regels, en uitgebreidere pagina's die een stuk groter zijn qua aantal regels en karakters. Het zijn dus waarschijnlijk de lengtes van deze uitgebreidere pagina's die log-normaal verdeeld zijn. Bij de derde figuur is een log-normale verdeling lastig waar te nemen. Dit komt waarschijnlijk doordat er veel korte regels van nul, een half of één woord(en) in de pagina's staan, die door een betere schoonmaak gefilterd zouden kunnen worden. Ook is het hier lastiger om een log-normale verdeling waar te nemen omdat de x-as van deze plot niet op logaritmische schaal is. 

3. Waarschijnlijk zijn er (na opschonen) heel veel pagina's op het internet te vinden waarbij het aantal karakters en regels van de pagina's log-normaal verdeeld zijn. 

# Deel 3: theorie

Deze vragen gaan over precisie, recall en accuraatheid.

Stel je hebt een ziekte die bij 1 op de 100 mensen voorkomt. En je hebt een test met een accuraatheid van .95. Dit laatste betekent dat je in 95% van de gevallen juist voorspelt of iemand de ziekte heeft of niet. Anders gezegd: de som van de elementen op de diagonaal van de confusion matrix gedeeld door het totaal aantal elementen is .95.

1. Wat is de maximale(variabele `maxR`) en minimale (`minR`) *recall* voor de ziekte-klasse bij deze gegevens?
2. Wat is de maximale (`maxP`) en minimale (`minP`) *precision* voor de ziekte-klasse bij deze gegevens?

**Hint** Maak een confusion matrix waarin je verschillende settings uitprobeert. Zorg er steeds voor dat alles netjes optelt. 

<!-- grade_id: d3a, points: 0, grade: false, solution: true -->
```python
data = {"has disease": ["100", "0", "100"],
        "has no disease": ["500", "9400", "9900"],
        "total": ["600", "9400", "10000"]}
df = pd.DataFrame(data, columns = ["has disease", "has no disease", "total"], 
                  index = ["test positive", "test negative", "total"])

data2 = {"has disease": ["0", "100", "100"],
        "has no disease": ["400", "9500", "9900"],
        "total": ["400", "9600", "10000"]}
df2 = pd.DataFrame(data2, columns = ["has disease", "has no disease", "total"], 
                  index = ["test positive", "test negative", "total"])

minP = 0 / 400 # precisie (minimale legitiem positieve tests t.o.v. alle positieve tests)
maxP = 100 / 600 # precisie (maximale legitiem positieve tests t.o.v. alle positieve tests)
minR = 0 / 100 # recall (minimale legitiem zieken t.o.v. alle ziek geclassificeerden)
maxR = 100 / 100 # recall (maximale legitiem zieken t.o.v. alle ziek geclassificeerden)

print('Minimale precisie: ', minP, '\nMaximale precisie: ', maxP, '\nMinimale recall: ', minR, '\nMaximale recall: ', maxR)
print('\nMaximale precisie en maximale recall: ')
display(df)
print('\nMinimale precisie en minimale recall: ')
display(df2)
```

<!-- grade_id: d3t, points: 1, grade: true, solution: false -->
```python
for m in {minP, maxP, minR, maxR}:
    assert 0 <= m <= 1
```

## Variant 

Dezelde vraag als boven, maar nu hebben 5 op de 100 mensen de ziekte en hebben we juist een accuraatheid van .99

<!-- grade_id: d31a, points: 0, grade: false, solution: true -->
```python
data3 = {"has disease": ["500", "0", "500"],
        "has no disease": ["100", "9400", "9500"],
        "total": ["600", "9400", "10000"]}
df3 = pd.DataFrame(data3, columns = ["has disease", "has no disease", "total"], 
                  index = ["test positive", "test negative", "total"])

data4 = {"has disease": ["400", "100", "500"],
        "has no disease": ["0", "9500", "9500"],
        "total": ["400", "9600", "10000"]}
df4 = pd.DataFrame(data4, columns = ["has disease", "has no disease", "total"], 
                  index = ["test positive", "test negative", "total"])

minP = 500 / 600 # precisie (minimale legitiem positieve tests t.o.v. alle positieve tests)
maxP = 400 / 400 # precisie (maximale legitiem positieve tests t.o.v. alle positieve tests)
minR = 400 / 500 # recall (minimale legitiem zieken t.o.v. alle ziek geclassificeerden)
maxR = 500 / 500 # recall (maximale legitiem zieken t.o.v. alle ziek geclassificeerden)

print('Minimale precisie: ', minP, '\nMaximale precisie: ', maxP, '\nMinimale recall: ', minR, '\nMaximale recall: ', maxR)
print('\nMinimale precisie en maximale recall: ')
display(df3)
print('\nMaximale precisie en minimale recall: ')
display(df4)
```

<!-- grade_id: d31t, points: 1, grade: true, solution: false -->
```python
for m in {minP, maxP, minR, maxR}:
    assert 0 <= m <= 1
```

# Deel 4 Corona: Wat is erger: besmettelijkheid of dodelijkheid?


Bron: <https://www.nrc.nl/nieuws/2020/12/31/verdachte-virusvariant-met-opvallend-veel-mutaties-a4025804>

>Wat is erger: besmettelijkheid of dodelijkheid? Op het eerste gezicht denk je het tweede. In de media is dan ook opluchting te bespeuren dat het virus weliswaar besmettelijker is, maar geen ernstiger ziekteverloop veroorzaakt. Maar is dat terecht? Adam Kucharski van de London School of Hygiene and Tropical Medicine [betoogde op Twitter](https://twitter.com/AdamJKucharski/status/1343567425107881986) van niet. Hij gaf het volgende getallenvoorbeeld:

>Stel, R is 1,1. Dat wil zeggen dat tien besmette personen gemiddeld elf anderen besmetten. Stel, de generatietijd is zes dagen: het duurt zes dagen tot die tien mensen elf anderen hebben besmet, dus in 30 dagen vertakt die ‘besmettingsboom’ zich vijf keer. Stel, van alle besmette mensen overlijdt 0,8 procent. En stel, op tijdstip nul zijn 10.000 mensen besmet. Dan levert dat na een maand $10.000 × (1,1)^5 × 0,8\% = 129$ doden op.

>Wat gebeurt er dan als het virus door een mutatie 50 procent dodelijker wordt? Dan zijn na een maand $10.000 x (1,1)^5 × (0,8\% × 1,5) = 193 $
mensen overleden.

>Als het virus door een mutatie echter 50 procent besmettelijker wordt, dan ligt het dodental bijna vijf keer zo hoog: $10.000 × (1,1 × 1,5)^5 × 0,8\% = 978$ doden.

>Hoe dat kan? De dodelijkheid correleert lineair met het aantal doden, maar voor besmettelijkheid is dat verband exponentieel.

* Run onderstaande code. Snap wat er gebeurt. Snap je het verschil in aantal doden voor periode 0. En waarom is het gelijk in periode 1, en gaat het daarna (steeds sneller) uit elkaar lopen?
* Check dat het aantal doden aan het eind van de maand in het artikel hetzelfde is als in de tabel bij periode 5.
* En kijk eens hoe makkelijk je een plotje van zo'n spreadsheet maakt. Ideaal toch? Later in de cursus leren we hier meer over.

## Klopt het wel? (1pt)

1. De formule klopt natuurlijk, maar de uitkomst niet. Want als je een maand met een R van 1.1 zit heb je 6 "generaties", en dus in totaal veel meer besmette mensen gehad dan $10.000 x (1,1)^5$.  **Hoeveel meer dan?** **Bereken dat in de functie `aantal_geinfecteerden_na_periodes`.**
2. En dus hebben we na een maand ook een ander aantal overledenen.  Bereken dat eens.



```python
# zoals in het artikel
def aantaldoden(aantalbesmet=10**4,
                overlijdingsratio=.008,
                R=1.1,
                hoeveelerger=1.5,
                aantalperiodes=5):
    return {
        'dodelijker':
        round(aantalbesmet * (R)**aantalperiodes * overlijdingsratio *
              hoeveelerger),
        'besmettelijker':
        round(aantalbesmet * (R * hoeveelerger)**aantalperiodes *
              overlijdingsratio)
    }


# 50 % erger
A = {P: aantaldoden(hoeveelerger=1.5, aantalperiodes=P) for P in range(0, 6)}
A = pd.DataFrame(A).T
A.index.name = 'Periode'
A.plot(kind='bar', title='Aantal doden van geinfecteerden in periode x.')
A
```

<!-- grade_id: CA, points: 0, grade: false, solution: true -->
```python
def aantal_geinfecteerden_na_periodes(
    aantalperiodes=6,
    aantalbesmet=10**4,
    R=1.1,
):
    '''Geef het totaal aantal geinfecteerde mensen terug na aantalperiodes generaties/periodes.
    Rond af op een geheel getal.'''
    # aantal geïnfecteerden, omdat 10000 mensen vervolgens 11000 anderen besmetten en er dus 21000 besmettingen zijn
    return round(sum([aantalbesmet * (R)**periode for periode in range(aantalperiodes + 1)]))

B = pd.Series({
    P: round(aantal_geinfecteerden_na_periodes(aantalperiodes=P))
    for P in range(21)
})
B.plot(kind='bar')
B
```

<!-- grade_id: CT, points: 1, grade: true, solution: false -->
```python
assert isinstance(aantal_geinfecteerden_na_periodes(), int)


```

## Aangepast dataframe (1pt)

Bekijk de code hierboven nog eens waarin het Dataframe `A` gemaakt werd. Pas die code aan, zodat je daarmee het dataframe `AA` maakt dat de juiste aantallen print. `AA` heeft dezelfde *shape* als `A`, en ook dezelfde index en kolom namen. Alleen de waardes zijn anders. 

* Apprecieer ook het verschil in de titels tussen de twee plotjes.

<!-- grade_id: C2A, points: 0, grade: false, solution: true -->
```python
# nieuw dataframe waarbij het aantal doden wordt opgeteld bij het aantal doden uit een rij eerder
AA = np.cumsum(A, axis=0)

T = 'Totaal aantal doden na x periodes van 6 dagen.'
AA.plot(kind='bar', title=T)
AA
```

<!-- grade_id: C2T, points: 1, grade: true, solution: false -->
```python
assert_equal(AA.shape, (6, 2))  # 2 kolommen, 6 rijen
assert_equal(list(AA.index) == list(range(6)),
             True)  # de index van AA is [0,1,2,3,4,5]


```

```python

```